# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # YOUR CODE HERE
    pattern = "[A-Z][\w]*"
    return re.findall(pattern,simple_string)
    
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

```python
#pattern = "[A-Z][\w]*"
#simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
#    Ruth and Peter, their parents, have 3 kids."""
```

```python
#re.findall(pattern, simple_string)
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

```python
#names()
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    # YOUR CODE HERE
    name_pattern = "(.*):\sB"
    return re.findall(name_pattern, grades)
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

```python
#with open ("assets/grades.txt", "r") as file:
#        grades = file.read()
#grades
```

Output:
```
'Ronald Mayr: A\nBell Kassulke: B\nJacqueline Rupp: A \nAlexander Zeller: C\nValentina Denk: C \nSimon Loidl: B \nElias Jovanovic: B \nStefanie Weninger: A \nFabian Peer: C \nHakim Botros: B\nEmilie Lorentsen: B\nHerman Karlsen: C\nNathalie Delacruz: C\nCasey Hartman: C\nLily Walker : A\nGerard Wang: C\nTony Mcdowell: C\nJake Wood: B\nFatemeh Akhtar: B\nKim Weston: B\nNicholas Beatty: A\nKirsten Williams: C\nVaishali Surana: C\nCoby Mccormack: C\nYasmin Dar: B\nRomy Donnelly: A\nViswamitra Upandhye: B\nKendrick Hilpert: A\nKillian Kaufman: B\nElwood Page: B\nMukti Patel: A\nEmily Lesch: C\nElodie Booker: B\nJedd Kim: A\nAnnabel Davies: A\nAdnan Chen: B\nJonathan Berg: C\nHank Spinka: B\nAgnes Schneider: C\nKimberly Green: A\nLola-Rose Coates: C\nRose Christiansen: C\nShirley Hintz: C\nHannah Bayer: B'
```

```python
# I've worked out how to find the grade ...
# this seem to find the colon with :
# then find the blank space with \s
# then get everything after these with (.*)
```

```python
#grade_pattern = ":\s(.*)"

#re.findall(grade_pattern, grades)
```

Output:
```
['A',
 'B',
 'A ',
 'C',
 'C ',
 'B ',
 'B ',
 'A ',
 'C ',
 'B',
 'B',
 'C',
 'C',
 'C',
 'A',
 'C',
 'C',
 'B',
 'B',
 'B',
 'A',
 'C',
 'C',
 'C',
 'B',
 'A',
 'B',
 'A',
 'B',
 'B',
 'A',
 'C',
 'B',
 'A',
 'A',
 'B',
 'C',
 'B',
 'C',
 'A',
 'C',
 'C',
 'C',
 'B']
```

```python
# Now how to find the name
# here we know that the colon is :
# then just ask for everything in front of that
# the full stop find you anything on a new line 
```

```python
#name_pattern = "(.*):"
#re.findall(name_pattern, grades)
```

Output:
```
['Ronald Mayr',
 'Bell Kassulke',
 'Jacqueline Rupp',
 'Alexander Zeller',
 'Valentina Denk',
 'Simon Loidl',
 'Elias Jovanovic',
 'Stefanie Weninger',
 'Fabian Peer',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Herman Karlsen',
 'Nathalie Delacruz',
 'Casey Hartman',
 'Lily Walker ',
 'Gerard Wang',
 'Tony Mcdowell',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Nicholas Beatty',
 'Kirsten Williams',
 'Vaishali Surana',
 'Coby Mccormack',
 'Yasmin Dar',
 'Romy Donnelly',
 'Viswamitra Upandhye',
 'Kendrick Hilpert',
 'Killian Kaufman',
 'Elwood Page',
 'Mukti Patel',
 'Emily Lesch',
 'Elodie Booker',
 'Jedd Kim',
 'Annabel Davies',
 'Adnan Chen',
 'Jonathan Berg',
 'Hank Spinka',
 'Agnes Schneider',
 'Kimberly Green',
 'Lola-Rose Coates',
 'Rose Christiansen',
 'Shirley Hintz',
 'Hannah Bayer']
```

```python
# I was trying to find a clever answer, but it's just to find the names to the left of ': B'
```

```python
#name_pattern = "(.*):\sB"
#re.findall(name_pattern, grades)
```

Output:
```
['Bell Kassulke',
 'Simon Loidl',
 'Elias Jovanovic',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Yasmin Dar',
 'Viswamitra Upandhye',
 'Killian Kaufman',
 'Elwood Page',
 'Elodie Booker',
 'Adnan Chen',
 'Hank Spinka',
 'Hannah Bayer']
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE
    pattern = """
    (?P<host>.*) #get everything 
    (\ -\ ) # up to the first space dash and space
    (?P<user_name>.*) # then get everything 
    (\s\[) # then there is a square bracket
    (?P<time>.*) # then get everything up to
    (\]\s\") # up to the square bracket, space and speech marks
    (?P<request>.*) # then everything 
    (") # then a final speech mark
    """
    # YOUR CODE HERE
    result = []
    for item in re.finditer(pattern, logdata, re.VERBOSE):
        result.append(item.groupdict())
    return result
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python
# Open the data
```

```python
#with open("assets/logdata.txt", "r") as file:
 #       logdata = file.read()
```

```python
#logdata[:10000]
```

```python
# lets just try and find the first set of numbers
# so we want everything prior to the first blank space and dash
```

```python
#first_pattern = "(.*) - "
#re.findall(first_pattern, logdata)
```

```python
#second_pattern = "-\s(.*)\s\["
#re.findall(second_pattern,logdata)
```

```python
# \\ makes sure that the regex can read the special characters, the square brackets
# then you're saying you want everything in between the two square brackets 
```

```python
#third_pattern = "\\[(.*)]"
#re.findall(third_pattern,logdata)
```

```python
#fourth_pattern = r'"(.*?)"'
#re.findall(fourth_pattern,logdata)
```

```python
"""
first_pattern = "(.*) - "
host_list = re.findall(first_pattern, logdata)
print(len(host_list))

second_pattern = "-\s(.*)\s\["
user_name_list = re.findall(second_pattern, logdata)
print(len(user_name_list))

third_pattern = "\\[(.*)]"
time_list = re.findall(third_pattern, logdata)
print(len(time_list))
    
fourth_pattern = r'"(.*?)"'
request_list = re.findall(fourth_pattern, logdata)
print(len(request_list))


# join these together into a generator
generator_of_dicts = (
    {"host": host, "user_name": user_name, "time": time, "request": request}
    for host, user_name, time, request in zip(host_list, user_name_list, time_list, request_list)
)

# convert this into a list
my_dictionary = list(generator_of_dicts)
my_dictionary
"""
```

Output:
```
'\nfirst_pattern = "(.*) - "\nhost_list = re.findall(first_pattern, logdata)\nprint(len(host_list))\n\nsecond_pattern = "-\\s(.*)\\s\\["\nuser_name_list = re.findall(second_pattern, logdata)\nprint(len(user_name_list))\n\nthird_pattern = "\\[(.*)]"\ntime_list = re.findall(third_pattern, logdata)\nprint(len(time_list))\n    \nfourth_pattern = r\'"(.*?)"\'\nrequest_list = re.findall(fourth_pattern, logdata)\nprint(len(request_list))\n\n\n# join these together into a generator\ngenerator_of_dicts = (\n    {"host": host, "user_name": user_name, "time": time, "request": request}\n    for host, user_name, time, request in zip(host_list, user_name_list, time_list, request_list)\n)\n\n# convert this into a list\nmy_dictionary = list(generator_of_dicts)\nmy_dictionary\n'
```

```python
"""
# both of these need to be correct to pass the test - which they are:
print(len(logs()) ==979)

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}

print(one_item in logs())
"""
```

Output:
```
"\n# both of these need to be correct to pass the test - which they are:\nprint(len(logs()) ==979)\n\none_item={'host': '146.204.224.152',\n  'user_name': 'feest6811',\n  'time': '21/Jun/2019:15:45:24 -0700',\n  'request': 'POST /incentivize HTTP/1.1'}\n\nprint(one_item in logs())\n"
```

```python
with open("assets/logdata.txt", "r") as file:
       logdata = file.read()
logdata[:1000]
```

Output:
```
'146.204.224.152 - feest6811 [21/Jun/2019:15:45:24 -0700] "POST /incentivize HTTP/1.1" 302 4622\n197.109.77.178 - kertzmann3129 [21/Jun/2019:15:45:25 -0700] "DELETE /virtual/solutions/target/web+services HTTP/2.0" 203 26554\n156.127.178.177 - okuneva5222 [21/Jun/2019:15:45:27 -0700] "DELETE /interactive/transparent/niches/revolutionize HTTP/1.1" 416 14701\n100.32.205.59 - ortiz8891 [21/Jun/2019:15:45:28 -0700] "PATCH /architectures HTTP/1.0" 204 6048\n168.95.156.240 - stark2413 [21/Jun/2019:15:45:31 -0700] "GET /engage HTTP/2.0" 201 9645\n71.172.239.195 - dooley1853 [21/Jun/2019:15:45:32 -0700] "PUT /cutting-edge HTTP/2.0" 406 24498\n180.95.121.94 - mohr6893 [21/Jun/2019:15:45:34 -0700] "PATCH /extensible/reinvent HTTP/1.1" 201 27330\n144.23.247.108 - auer7552 [21/Jun/2019:15:45:35 -0700] "POST /extensible/infrastructures/one-to-one/enterprise HTTP/1.1" 100 22921\n2.179.103.97 - lind8584 [21/Jun/2019:15:45:36 -0700] "POST /grow/front-end/e-commerce/robust HTTP/2.0" 304 14641\n241.114.1
```

```python
# My technique worked, but they use another way of getting the results
# I will try this to practice my regex



#pattern = """
#(?P<host>.*) #get everything 
#(\ -\ ) # up to the first space dash and space
#(?P<user_name>.*) # then get everything 
#(\s\[) # then there is a square bracket
#(?P<time>.*) # then get everything up to
#(\]\s\") # up to the square bracket, space and speech marks
#(?P<request>.*) # then everything 
#(") # then a final speech mark
#"""
#store = []
#for item in re.finditer(pattern, logdata, re.VERBOSE):
#    print(item.groupdict())
#    store.append(item)
```